Spatial integration of sensory input and motor output in Pseudomonas aeruginosa chemotaxis through colocalized distribution

Author response:

The following is the authors’ response to the previous reviews.

Reviewer #1 (Public review):

Summary:

The study by Wu et al presents interesting data on bacterial cell organization, a field that is progressing now, mainly due to the advances in microscopy. Based mainly on fluorescence microscopy images, the authors aim to demonstrate that the two structures that account for bacterial motility, the chemotaxis complex and the flagella, colocalize to the same pole in Pseudomonas aeruginosa cells and to expose the regulation underlying their spatial organization and functioning.

Comments on revisions:

The authors have addressed all major and minor points that I raised in a satisfying way during the revision process. The work can now be regarded as complete, the assumptions were clarified, the results are convincing, the conclusions are justified, and the novelty has been made clear.

This manuscript will be of interest to cell biologists, mainly those studying bacteria, but not only.

Reviewer #2 (Public review):

Summary:

Here, the authors studied the molecular mechanisms by which the chemoreceptor cluster and flagella motor of Pseudomonas aeruginosa (PA) are spatially organized in the cell. They argue that FlhF is involved in localizing the receptors-motor to the cell pole, and even without FlhF, the two are colocalized. Finally, the authors argue that the functional reason for this colocalization is to insulate chemotactic signaling from other signaling pathways, such as cyclic-di-GMP signaling.

Strength:

The experiments and data are high quality. It is clear that the motor and receptors co-localize, and that elevated CheY levels lead to elevated c-di-GMP.

Weakness:

The explanation for the functional importance of receptor-motor colocalization is plausible but is still not conclusively demonstrated. Colocalization might reduce CheY levels throughout the cell in order to reduce cross-talk with c-di-GMP. This would mean that if physiologically-relevant levels of CheYp near the pole were present throughout the cell, c-di-GMP levels would be elevated to a point that is problematic for the cell. Clearly demonstrating this seems challenging.

We acknowledge that directly proving the necessity of colocalization to prevent problematic c-di-GMP elevation is experimentally challenging, as it would require creating a system where CheY-P is artificially distributed throughout the cell at physiologically relevant concentrations while maintaining normal chemotaxis function.

However, our data provide several lines of evidence supporting this model. First, we show that CheY overexpression leads to substantial c-di-GMP elevation (71.8% increase) and cell aggregation, demonstrating that elevated CheY levels can indeed cause problematic cross-pathway interference. Second, previous work has shown that CheY-P levels near the pole are an order of magnitude higher than in the rest of the cell (ref. 46). If this elevated CheY-P concentration near the pole were present throughout the cell, our data suggest that c-di-GMP levels would be elevated sufficiently to cause cell aggregation (Fig. 4A), thereby disabling normal motility and chemotaxis. Third, the dose-dependent relationship between CheY concentration and aggregation phenotype supports the idea that precise spatial regulation of CheY levels is functionally important for avoiding cross-pathway interference.

Reviewer #3 (Public review):

Summary:

The authors investigated the assembly and polar localization of the chemosensory cluster in P. aeruginosa. They discovered that a certain protein (FlhF) is required for the polar localization of the chemosensory cluster while a fully-assembled motor is necessary for the assembly of the cluster. They found that flagella and chemosensory clusters always co-localize in the cell; either at the cell pole in wild type cells or randomly-located in the cell in FlhF mutant cells. They hypothesize that this co-localization is required to keep the level of another protein (CheY-P), which controls motor switching, at low levels as the presence of high-levels of this protein (if the flagella and chemosensory clusters were not co-localized) is associated with high-levels of c-di-GMP and cell aggregations.

Strengths:

The manuscript is clearly written and straightforward. The authors applied multiple techniques to study the bacterial motility system including fluorescence light microscopy and gene editing. In general, the work enhances our understanding of the subtlety of interaction between the chemosensory cluster and the flagellar motor to regulate cell motility.

Weaknesses:

The major weakness for me in this paper is that the authors never discussed how the flagellar genes expression is controlled in P. aeruginosa. For example, in E. coli there is a transcriptional hierarchy for the flagellar genes (early, middle, and late genes, see Chilcott and Hughes, 2000). Similarly, Campylobacter and Helicobacter have a different regulatory cascade for their flagellar genes (See Lertsethtakarn, Ottemann, and Hendrixson, 2011). How does the expression of flagellar genes in P. aeruginosa compare to other species? how many classes are there for these genes? is there a hierarchy in their expression and how does this affect the results of the FliF and FliG mutants? In other words, if FliF and FliG are in class I (as in E. coli) then their absence might affect the expression of other later flagellar genes in subsequent classes (i.e., chemosensory genes). Also, in both FliF and FliG mutants no assembly intermediates of the flagellar motor are present in the cell as FliG is required for the assembly of FliF (see Hiroyuki Terashima et al. 2020, Kaplan et al. 2019, Kaplan et al. 2022). It could be argued that when the motor is not assembled then this will affect the expression of the other genes (e.g., those of the chemosensory cluster) which might play a role in the decreased level of chemosensory clusters the authors find in these mutants.

We thank the reviewer for the valuable suggestions. In the revised manuscript, we have further elaborated on the regulatory control of flagellar genes expression in P. aeruginosa (see our response to comment #4).

Comments on revisions:

I believe the authors have performed additional experiments that improved their manuscript and they have answered many of my comments and those of the other reviewers. I am supportive of publishing this manuscript, but I still find the following points that are not clear to me (probably I am misunderstanding some points; the authors can clarify).

(1) In response to reviewer 1, the authors say that they "analyzed and categorized the distribution of the chemotaxis complex in both wild-type and flhF mutant strains into three patterns: precise-polar, near-polar, and mid-cell localization." I can see what they mean by polar and mid-cell, but near-polar sounds a bit elusive? Can they provide examples of this stage and mention how accurately they can identify it? Also, do the pie charts they show in Figure S4 really show "significant alterations"? There is a difference between 98% and 85% as they mention in their response to reviewer 1, but I am not sure that this is significant? Probably they can explain/change the language in the text? Also, the number of cells they counted for FlhF mutant is more than the double of other strains (WT and FlhF FliF mutant)?

We thank the reviewer for the valuable suggestions. To clarify, we divided the intracellular area along the cell's long axis into three domains: the two ends each representing 10% of the length as the precise-polar domain, the central 50% as the mid-cell domain, and the remaining regions between these as the near-polar domain. The localization pattern of the chemotaxis complex was assigned based on the position of the fluorescence intensity centroid within these domains.

Regarding the significance of the changes, you are correct to question our language. When flhF was knocked out, the proportion of chemotaxis complexes with precise-polar distribution decreased from 98% to 85% - a 13% reduction. While this represents a measurable shift in localization pattern, describing this as "significant alterations" was probably imprecise. We have revised this language to more accurately reflect the magnitude of the change (lines 169-177).

For the cell counting, we increased the sample size for the flhF mutant because this strain exhibited the appearance of mid-cell localization (approximately 5% of cells), which was not observed in wild-type or flhF fliF double mutant strains. To accurately quantify this rare phenotype and ensure statistical reliability, we analyzed more cells for this particular strain. This explains why the flhF mutant dataset contains approximately double the number of cells compared to the other strains.

We have redrawn Figure S4 to include a clear schematic diagram of the cell partitioning method and provided representative examples of each localization pattern (precise-polar, near-polar, and mid-cell) to better illustrate how we distinguished between these categories.

(2) One thing that also confused me is the following: One point that the authors stress is that FlhF localizes both the flagellum and the chemoreceptors to the pole. However, if I look at Figure 2B, the flagellum and the chemoreceptors still co-localize together (although not at the pole). If FlhF was responsible for co-localizing both of them to the pole, then wouldn't one expect them to be randomly localized in this mutant and by that I mean that they do not co-localize but that each of them (the flagellum and the chemoreceptors) are located in a different random location of the cell (not co-localized). The fact that they are still co-localized together in this mutant could also be interpreted by, for example, that FlhF localizes the flagellum to the pole and another mechanism localizes the chemoreceptors to the flagellum, hence, they still co-localize in this mutant because the chemoreceptors follow the flagellum by another mechanism to wherever it goes?

Thank you for this insightful observation. You are correct that our current experimental results do not definitively establish that FlhF directly localizes both the flagellum and chemoreceptors to the pole independently. The persistent colocalization of flagella and chemoreceptors in the DflhF mutant, even when both are mislocalized away from the pole, actually suggests a more complex regulatory mechanism than we initially proposed.

This observation highlights an important distinction between polar targeting and colocalization maintenance. Our data suggest that FlhF influences the polar targeting of the flagellum-chemoreceptor assembly, but the colocalization itself appears to be governed by a different mechanism that operates independently of FlhF. This could involve direct protein-protein interactions between flagellar and chemotaxis components, or shared assembly machinery that we have yet to identify.

To better reflect this interpretation, we have revised the subsection title (line 150). We have also modified the relevant discussion (line 180) to more accurately describe FlhF’s role in polar targeting rather than claiming it directly controls chemoreceptor localization.

(3) In the response to reviewers, the authors mention "suggesting that the assembly of the receptor complex is likely influenced mainly by the C-ring and MS-ring structures rather than by the P ring". However, in the article, they still write "The complete assembly of the motor serves as a partial prerequisite for the assembly of the chemotaxis complex, and its assembly site is also regulated by the polar anchor protein FlhF" despite their FlgI results which is not in accordance with this statement? Also, As I mentioned in my previous report, in FliG and FliF mutant the motor does not assemble (see Hiroyuki Terashima et al. 2020., and Kaplan et al., 2022).

We thank the reviewer for the suggestions and acknowledge the contradictions in our original text. You are correct that in DfliF and DfliG mutants, the flagellar motor does not assemble, while the P ring (FlgI) functions as a bushing for the peptidoglycan layer and its absence does not prevent motor assembly.

Our DflgI results, which showed normal chemotaxis complex assembly similar to wild-type, clearly demonstrate that the P ring is not required for chemoreceptor complex formation. This contradicts our original statement that "complete assembly of the motor serves as a partial prerequisite for the assembly of the chemotaxis complex."

We have corrected this inconsistency by: 1) Revising the subsection title (line 186) to more accurately reflect that core motor structures, rather than complete motor assembly, influences chemoreceptor complex formation. 2) Modifying sentences in the introduction (lines 97-98) to better align with our experimental findings.

(4) The authors have said in their response to my point "and currently, there is no evidence that FliA activity is influenced by proteins like FliG". I just want to clarify what I meant in my previous report: In E. coli, FliA binds to FlgM, and when the hook is assembled FlgM is secreted outside the cell allowing FliA to trigger the transcription of class III genes, which include the chemosensory genes (see Figure 5 in Beeby et al, 2020 in FEMS Microbiology, and Chilcott and Hughes, 2000). This implies that if the hook is not built, then late genes (including the chemoreceptors) should not be present. However, in Kaplan et al., 2019, the authors imaged a FliF mutant in Shewanella oneidensis (Figure S3) and still saw that chemoreceptors are present (I believe the authors must highlight this). This suggests that species such as Shewanella and Pseudomonas have a different assembly process than that E. coli, and although the authors say that in the text, I believe they still can refine this part more in the spirit of what I wrote here.

We thank the reviewer for the important clarification regarding the differences in transcriptional regulation among bacterial species. We agree that the observation of chemoreceptors in Shewanella oneidensis DfliF mutants (Kaplan et al., 2019) represents a significant deviation from the well-characterized E. coli model and merits stronger emphasis. In response, we have expanded the discussion to more clearly highlight the critical distinctions in the transcriptional regulatory circuits governing flagellar and chemoreceptor biogenesis between E. coli and species such as Shewanella oneidensis and Pseudomonas aeruginosa (lines 351-363).

I do not like to ask for additional experiments in the second round of review, so for me if the authors modify the text to tackle these points and allow for probable alternative explanations/ highlight gaps/ modify language used for some claims, then that is fine with me.

Reviewer #2 (Recommendations for the authors):

It is plausible that colocalization reduces CheY levels throughout the cell in order to reduce cross-talk with c-di-GMP. This would mean that if physiologically-relevant levels of CheYp near the pole were present throughout the cell, c-di-GMP levels would be elevated to a point that is problematic for the cell. Clearly demonstrating this seems challenging.

We acknowledge that directly proving the necessity of colocalization to prevent problematic c-di-GMP elevation is experimentally challenging, as it would require creating a system where CheY-P is artificially distributed throughout the cell at physiologically relevant concentrations while maintaining normal chemotaxis function.

However, our data provide several lines of evidence supporting this model. First, we show that CheY overexpression leads to substantial c-di-GMP elevation (71.8% increase) and cell aggregation, demonstrating that elevated CheY levels can indeed cause problematic cross-pathway interference. Second, previous work has shown that CheY-P levels near the pole are an order of magnitude higher than in the rest of the cell (ref. 46). If this elevated CheY-P concentration near the pole were present throughout the cell, our data suggest that c-di-GMP levels would be elevated sufficiently to cause cell aggregation (Fig. 4A), thereby disabling normal motility and chemotaxis. Third, the dose-dependent relationship between CheY concentration and aggregation phenotype supports the idea that precise spatial regulation of CheY levels is functionally important for avoiding cross-pathway interference.

Changes from the earlier version to the revised version

Abstract

Abstract

The opportunistic pathogen Pseudomonas aeruginosa serves as a model organism for studying multiple signal transduction pathways. The chemoreceptor cluster, a core component of the chemotaxis pathway, is assembled from hundreds of proteins. The unipolar distribution of receptor clusters has long been recognized, yet the precise mechanism governing their assembly remains elusive. Here, we directly observed the relative positions of the flagellar motor and chemoreceptor cluster using flagellar filament labeling and gene editing techniques. Surprisingly, we found that both are located at the same cell pole, with the distribution pattern controlled by the polar anchor protein FlhF. Additionally, the efficient assembly of the chemoreceptor cluster is partially dependent on the integrity of the motor structure. Furthermore, we discovered that overexpression of the chemotaxis regulatory protein CheY leads to high intracellular levels of the second messenger c-di-GMP, triggering cell aggregation. Therefore, the colocalization of the chemoreceptor cluster and flagellum in P. aeruginosa serves to avoid cross-pathway signaling interference, enabling cells to conduct various physiological activities in an orderly manner.

Introduction

Earlier version: In biological systems, the process by which proteins self-assemble into organized complex structures is widespread1,2. Pole-to-pole protein oscillations in the Min system ensure that the FtsZ ring, a crucial component of cell division, is placed precisely in the middle of the cell body3,4. In response to antibiotic treatment and heat stress, some bacteria generate multiple protein foci throughout the cytoplasm5,6. The polar localization of the secretion system, such as the type VI secretion system, is mediated by targeting proteins and potentially facilitates host-pathogen interactions7,8. Thus, spontaneous assembly is accompanied by the formation of spatial patterns that organize the intracellular environment.

Revised version: In biological systems, the process by which proteins self-assemble into organized complex structures is widespread1,2. Pole-to-pole protein oscillations in the Min system ensure that the FtsZ ring, a crucial component of cell division, is placed precisely in the middle of the cell body3,4. In response to antibiotic treatment and heat stress, some bacteria generate multiple protein foci throughout the cytoplasm5,6. The polar localization of the secretion system, such as the type [secretion system, is mediated by targeting proteins and potentially facilitates host-pathogen interactions7,8. Thus, spontaneous assembly is accompanied by the formation of spatial patterns that organize the intracellular environment.

To grow and colonize in various ecological niches, bacteria have evolved a mechanism for migrating towards more favorable environments. The locomotion organ (the flagellar motor) and the chemotaxis pathway play crucial roles in achieving this goal. The former provides physical drive, while the latter offers directional guidance. The initiation of cell reorientation is controlled by motor switching, which is governed by the chemotaxis two-component signaling pathway. This pathway comprises transmembrane chemoreceptors, a histidine kinase CheA, an adaptor protein CheW, a response regulator CheY, and two adaptation proteins, CheB and CheR. Receptors collaborate with CheA and CheW to form chemotaxis complex in a hexagonal pattern9,10. Upon phosphorylation of CheY by the kinase at the complex, it freely diffuses to the motor and binds to FliM, a component of the motor switch complex, to affect motor switching11–13. It is worth noting that there exists diversity among different bacterial species in the placement of flagellar motors14, whereas chemoreceptor clusters are typically found at one or both ends of the cell15. These two structural units are functionally connected. However, it remains unclear whether there is an interaction in their spatial distribution and what the specific regulatory mechanism might be.

Pseudomonas aeruginosa, a common human opportunistic pathogen, possesses four chemosensory pathways that perform distinct functions and are stimulated by signal binding to 26 chemoreceptors16. Among them, the proteins encoded by chemotaxis-related genes collectively constitute the F6 pathway. Previous studies have suggested that the receptors involved in this pathway localize to the old cell pole17, similar to the flagellum18, implying colocalization of the receptors and the flagellum at the same pole; however, direct evidence for this colocalization is still lacking. Furthermore, the polar anchor protein FlhF has been reported to guide the flagellum to grow at the cell pole19,20. However, the mechanism by which chemotaxis complex-related proteins aggregate and self-assemble at the cell pole is controversial. The distribution of the chemotaxis complex in the peritrichous bacterium Escherichia coli has been extensively studied and attributed to several mechanisms, including stochastic self-assembly21,22, membrane curvature sorting23,24, and inefficient clustering in the lateral region25. However, these mechanisms cannot explain the unipolar distribution pattern observed in P. aeruginosa. Unlike E. coli, specific genes responsible for the placement of the chemotaxis complex have been identified in several bacterial species. In Caulobacter crescentus, chemotaxis complex assemble at the new cell pole with the help of TipN and TipF proteins26. A tripartite ParC-ParP-CheA interaction network was reported to promote polar localization of chemotaxis complex in Vibrio parahaemolyticus27,28. However, as there are no related genes in P. aeruginosa, the mechanism of its chemotaxis complex distribution remains a mystery.

Earlier version: Here, we combined gene editing and in vivo fluorescence imaging of flagellar filaments to directly observe the distribution of chemotaxis complex and flagellar motors, proposing a cooperative construction model of chemotaxis network and flagella during the entire division cycle of P. aeruginosa. The core focus of this study is to clarify the regulatory mechanism of its chemotaxis complex distribution. We found a substantial association between the assembly of the flagellar motor and the chemotaxis complex. The complete assembly of the motor serves as a partial prerequisite for the assembly of the chemotaxis complex, and its assembly site is also regulated by the polar anchor protein FlhF. Furthermore, by introducing exogenously expressed CheY protein, we found that this triggers the expression of c-di-GMP at a high level. From this, we infer that the colocalization of the chemotaxis complex and the flagellum in P. aeruginosa avoids the cross-pathway interference of signaling molecules, thus providing a guarantee for the coexistence of multiple chemosensory pathways.

Revised version: Here, we combined gene editing and in vivo fluorescence imaging of flagellar filaments to directly observe the distribution of chemotaxis complex and flagellar motors, proposing a cooperative construction model of chemotaxis network and flagella during the entire division cycle of P. aeruginosa. The core focus of this study is to clarify the regulatory mechanism of its chemotaxis complex distribution. We found a substantial association between the assembly of the flagellar motor and the chemotaxis complex. The assembly efficiency of the receptor complex is influenced by core flagellar motor components, particularly the C-ring and MS-ring structures, while its assembly site is also affected by the polar anchor protein FlhF. Furthermore, by introducing exogenously expressed CheY protein, we found that this triggers the expression of c-di-GMP at a high level. From this, we infer that the colocalization of the chemotaxis complex and the flagellum in P. aeruginosa avoids the cross-pathway interference of signaling molecules, thus providing a guarantee for the coexistence of multiple chemosensory pathways.

Results

Robust generation of daughter cell with both chemotaxis network and flagellar motor

CheY has been shown to colocalize with chemoreceptors17,29. To visualize the distribution of chemotaxis complex in cells, we fused the gene encoding the enhanced yellow fluorescent protein (eYFP) to the cheY gene in the P. aeruginosa chromosome (Fig. 1A). All genetic modifications were carried out at the native position on the chromosome, allowing the expression of chemotaxis proteins in precise stoichiometry from their natural promoters (Fig. 1B). The mutant strain with cheY-eyfp was able to moderately expand on soft agar plates17, proving that its chemotaxis behavior was not notably affected, so it was selected for subsequent experiments. As shown in Fig. 1C, CheY-eYFP was mainly located at the cell pole, suggesting that CheA, CheW and CheY in P. aeruginosa form a signal transduction complex that was primarily distributed at the cell pole, as described previously17. Representative large-field images containing many cells are shown in Fig. S1. Additionally, we constructed a plasmid expressing CheA-CFP and electroporated it into the cheY-eyfp strain. Fluorescence imaging revealed a high degree of spatial overlap between CheA-CFP and CheY-EYFP (Fig. S2), confirming that CheY-EYFP accurately marks the location of the chemoreceptor complex. From our measurements, nearly 90% (332/372) of cells contain obvious receptor clusters, and factors such as fluorescence bleaching may have caused the loss of fluorescent spots in some cells. In addition, we observed the presence of fluorescent spots at both ends of some cells with a large aspect ratio (probably approaching cell division), indicating that the newly generated progeny cells will have a complete chemotaxis network. Next, we sought to observe the distribution of the flagella in cells simultaneously to understand the motility of future progeny.

A. Schematic diagram of the chemotaxis signal transduction network and flagellar motor of P. aeruginosa. Fusion protein CheY-EYFP and fluorescently labeled flagellar filaments were used as markers to indicate the position of the chemotaxis complex and motor, respectively. B. The labeling mechanism of flagellar filaments and chemotaxis regulatory protein CheY. Filaments (with cysteine point mutation FliCT394C) were labeled through sulfhydryl-maleimide conjugation, and cheY-eyfp fusion with a 3× glycine linker was used to visualize chemotaxis complex positions. C. Localization of CheY-EYFP in the wild-type strain of P. aeruginosa. CheY-EYFP is mainly located at the single cell pole, and the white arrow points to individuals with obvious chemotaxis complex at both cell poles, which generally have a large aspect ratio of the cell body. The yellow dashed box marks the cell outline. D. The merged imaging of flagellar filaments and CheY-EYFP in the wild-type strain of P. aeruginosa, where flagellar motor and chemotaxis complex colocalize in cells. 145 cells with labeled flagella were observed, all of which exhibited consistent colocalization. White arrows point to individuals about to be divided, and the yellow dashed box marks the cell outline. The scale bar is 1 μm.

Our recent study demonstrated the successful labeling of the flagellar filament with thiol-reactive fluorescent dye by introducing cysteines into the flagellin FliC30,31 (Fig. 1B). To avoid fluorescence interference, we utilized a dye with an excitation peak near 568 nm, considering that the excitation peak of eYFP is 513 nm, and the emission peak is 527 nm. We employed a xenon lamp as the excitation light source and switched filters to enable simultaneous fluorescence imaging of flagellar filaments and chemotaxis complex within the same cell.

Remarkably, our findings revealed a surprising colocalization of chemotaxis complex and flagellar filaments at the same end of the cell body, and this colocalization was consistent, as shown in Fig. 1D. A typical large field was shown in Fig. S3. Similar to the chemotaxis complex, we observed the phenomenon of double flagella symbiosis in cells with a large aspect ratio, suggesting that P. aeruginosa has evolved a cell division mechanism with precise timing regulation. Before cell division, both poles of the mother cell assembled chemotaxis complex and flagellar motors. This ensures that daughter cells possess complete motility and chemotaxis, thereby greatly enhancing the environmental adaptability of the population.

Added in the revised version: FlhF controls polar targeting but not flagellum-chemoreceptor colocalization

Removed in the revised version: FlhF mediates proper unipolar localization of the chemotaxis complex

Despite potential differences in the physical and especially physiological environments at the two cell poles, it is unlikely that the unipolar distribution of the chemotaxis complex can be attributed to passive regulatory factors. The number and distribution of flagellar motors and chemotaxis complex vary among different bacterial species14,15, and the localization system for flagella has been well studied. Specifically, the FlhF-FlhG system, discovered in several monotrichous bacterial species, has been shown to control the location and number of flagella32,33. In P. aeruginosa, a knockout of flhF leads to mis-localized flagellar assembly19,20. Considering the consistent colocalization pattern between chemotaxis complex and flagellar motors in P. aeruginosa, we speculate that the distribution of chemotaxis complex is also regulated by similar molecular mechanisms.

Earlier version: To investigate the role of FlhF in the localization of chemotaxis complex, we constructed a ΔflhF strain for fluorescence observation. The results revealed that the chemotaxis complex no longer grow robustly at the cell pole (Fig. 2A), and the assembly positions of the flagellar motor change accordingly, colocalizing with the chemotaxis complex (Fig. 2B). Representative large fields containing many cells are shown in Fig. S1. Additionally, we categorized receptor cluster distribution patterns into three types: precise polar, near polar, and mid-cell. The relevant statistics for the wild-type strain and the ΔflhF mutant are presented in Fig. S4, which show a significant increase in the proportion of non-precise polar chemoreceptors following flhF knockout. We also quantified the proportion of individuals with obvious receptor clusters, which was more than 80% (182/221), similar to the wild-type strain. These experimental results suggest that the polar anchoring protein FlhF has a minimal impact on the assembly efficiency of P. aeruginosa’s chemotaxis complex, but it strictly controls its unipolar distribution. Furthermore, a consistent colocalization between the flagellar motor and chemotaxis complex was observed, independent of FlhF. Thus, it is crucial to determine whether there is a causal relationship in the assembly order of these two structural units.

Revised version: To investigate the role of FlhF in the localization of chemotaxis complex, we constructed a ΔflhF strain for fluorescence observation. The results revealed that the chemotaxis complex no longer grow robustly at the cell pole (Fig. 2A), and the assembly positions of the flagellar motor change accordingly, colocalizing with the chemotaxis complex (Fig. 2B). Representative large fields containing many cells are shown in Fig. S1. Additionally, we categorized receptor cluster distribution patterns into three types: precise polar, near polar, and mid-cell. We artificially divided the intracellular area along the long axis of the cell. The two ends of the length accounted for 10% each, which we called the precise-polar domains, the middle 50% became the mid-cell domain, and the rest was called the near-polar domain. The type of complex distribution was determined by the partition where the center of the chemotactic complex fluorescent bright spot was located. The relevant statistics for the wild-type strain and the ΔflhF mutant are presented in Fig. S4, showing that the proportion of precise polar distribution of chemotactic complexes decreased by 12.8% after flhF knockout. We also quantified the proportion of individuals with obvious receptor clusters, which was more than 80% (182/221), similar to the wild-type strain. These experimental results suggest that the polar anchoring protein FlhF has a minimal impact on the assembly efficiency of P. aeruginosa’s chemotaxis complex, but it affects its unipolar distribution. Furthermore, a consistent colocalization between the flagellar motor and chemotaxis complex was observed, independent of FlhF. Thus, it is crucial to determine whether there is a causal relationship in the assembly order of these two structural units.

A. Localization of CheY-EYFP in the ΔflhF strain of P. aeruginosa. CheY-EYFP is no longer robustly distributed at the single cell pole. The yellow dashed box marks the cell outline. B. The merged imaging of flagellar filaments and CheY-EYFP in the ΔflhF strain of P. aeruginosa, flagellar motor and chemotaxis complex still colocalize in cells. 101 cells with labeled flagella were observed, all of which exhibited consistent colocalization. The yellow dashed box marks the cell outline. The scale bar is 1 μm.

Added in the revised version: Core motor structures influence chemoreceptor complex assembly

Removed in the revised version: Motor structural integrity is a prerequisite for chemoreceptor self-assembly

Cryo-electron microscopy has successfully revealed the complete flagella structure in various bacterial species. This structure includes multiple components such as the inner-membrane MS ring, the cytoplasmic C ring, and the internal protein secretion system, all of which exhibit strong similarities34. This suggests that the core structure of the flagellar motor is highly conserved. The MS ring, composed of 26 FliF proteins, forms the foundation of the flagellar structure35. In addition to serving as a mounting platform for the C ring, the MS ring also acts as a protective shell for the internal protein secretion machinery36. FliG, the main component protein of the C ring, binds directly to the cytoplasmic surface of FliF in a 1:1 ratio37. Thus, the efficient assembly of the C ring requires preassembly of the MS ring. To observe the distribution of the chemotaxis complex following the disruption of the C ring and MS ring, we constructed ΔfliG and ΔfliF mutants.

We performed fluorescence observation on the ΔfliG mutant, which has a disrupted flagellar motor C ring. The results showed that the proportion of ΔfliG cells with obvious chemoreceptor clusters decreased significantly compared to the wild-type cells (59.8%, 140/234), although the chemotaxis complex remained at a single cell pole. We also examined the ΔfliF mutant, which has a disrupted flagellar motor MS ring. Similar to the ΔfliG mutant, the proportion of individuals with obvious chemotaxis complex decreased significantly (62.5%, 202/323). Subsequently, we constructed a ΔflhFΔfliF mutant and observed that the proportion of individuals with obvious chemoreceptor clusters further decreased (50.7%, 341/672). To ascertain whether it is motor integrity rather than functionality that influences the efficiency of chemotaxis complex assembly, we constructed a stator mutant (ΔmotAΔmotCD). In this mutant, the motor is completely stalled while the structure remains intact. We found that the mutant performed similarly to the wild-type strain in terms of chemotaxis complex assembly (84.3% of individuals with obvious clusters, 204/242). The fluorescence imaging of receptor clusters for multiple mutants in this section is shown in Figure. 3A, and corresponding large-field images are provided in Fig. S1. We utilized Western blotting to measure the expression levels of CheY, which were found to be similar across the different strains (Fig. 3B). These further substantiated that the observed phenomenon is based on the structural integrity of the motor rather than the protein expression level. In contrast, the cheA knockout strain was constructed to disrupt the assembly of the chemoreceptor complex. We fluorescently labeled its flagellar filaments and found that its phenotype was no different from that of the wild-type strain (Fig. S5), indicating that the assembly of the chemotactic receptor complex does not affect the flagellar assembly efficiency. The P ring (mainly composed of FlgI) is thought to act as a bushing for the peptidoglycan layer, and its absence results in partial assembly of the motor structure38. We constructed ΔflgI mutant and found that the proportion of cells with distinct chemotactic complexes was similar to that of the wild-type (Fig. S6). Additionally, we quantified the proportion of cells with receptor clusters (Fig. 3C). Overall, our findings suggest that the polar anchor protein FlhF and the structural integrity of the motor are crucial for the formation of chemotaxis complex. The former guides them to the appropriate site, while the latter influences the assembly efficiency of the chemoreceptor complex.

A. Localization of CheY-EYFP in various P. aeruginosa strains. The scale bar is 10 μm. B. Western blot analysis was performed to detect CheY expression in various P. aeruginosa strains. β-actin was used as the housekeeping protein. C. Occurrence probability of obvious chemotaxis complex in P. aeruginosa wild-type and several mutant strains. The proportion of individuals with obvious chemotaxis complex decreased significantly in the motor-incomplete strains (ΔfliF and ΔfliG), and this value was further reduced after flhF knockout (ΔfliFΔflhF). The ΔfliF and ΔmotAΔmotCD strains have a similar chemotaxis complex occurrence probability as the wild-type strain. The number of cells analyzed for each strain (from left to right) were: 372, 221, 234, 323, 672, and 242. ‘***’: significant difference (P-value < 0.0001), ‘ns.’: no significant difference (P-value > 0.05).

Colocalization of chemotaxis complex and flagellar motor avoids cross-pathway interference in P. aeruginosa signal transduction

The distribution patterns of receptors associated with various signal transduction pathways in P. aeruginosa vary considerably. For instance, the biofilm formation-related receptor WspA, unlike the chemotaxis receptors discussed here, is distributed throughout the cell. This distribution is thought to enhance its sensitivity to mechanical perturbations of the cell membrane39. Previous studies speculated that the highly consistent position of chemotaxis complex and flagellar motors enhances bacterial chemotaxis performance28. However, the diffusion time of the phosphorylated chemotaxis regulatory protein CheY-P across the longest distance in a bacterial cell body (along the cell’s long axis) is approximately 100 ms40,41, whereas the timescale for the chemotaxis temporal comparison is on the order of seconds42. Additionally, a study by Fukuoka and colleagues reported that intracellular chemotaxis signal transduction requires approximately 240 ms beyond CheY or CheY-P diffusion time41. Moreover, the CCW/CW interval of the P. aeruginosa flagellar motor under normal conditions is 1-2 s, as determined by bead assays30 or tethered cell assays43. Taken together, these indicate that for P. aeruginosa, which moves via a run-reverse mode, the potential 100 ms reduction in response time due to co-localization of the chemotaxis complex and motor has a limited effect on overall chemotaxis timing. Consequently, the physiological significance of the colocalization of chemoreceptors and flagellar motors remains unclear.

E. coli mediates chemotaxis through a single pathway involving five types of chemoreceptors44. As a core regulatory protein for chemotaxis, CheY participates in multiple processes such as adaptation and chemotaxis response, with CheY-P binds to the motor to regulate switching11. In contrast, P. aeruginosa has four distinct chemosensory pathways that perform different functions, involving different CheY homologs, and are stimulated by signals binding to 26 types of chemoreceptors. Consequently, P. aeruginosa possesses more complex chemosensory pathways16. We thus hypothesized that the co-localization of chemoreceptors and flagellar motors in P. aeruginosa ensures locally distributed CheY-P molecules, thereby eliminating the need for a high level of intracellular CheY-P and avoiding potential side effects on other signaling pathways.

To test the potential effect of an increased intracellular CheY-P level, we constructed the CheY expression plasmid cheY-pJN105, transformed it into the wild-type P. aeruginosa strain, and induced it with varying concentrations of arabinose. We observed that the higher the inducer concentration, the more pronounced the cell aggregation became in the field of view (Fig. 4A). The transition from planktonic individuals to aggregated communities is similar to biofilm formation, which is known to be accompanied by a significant increase in intracellular c-di-GMP levels.

A. The evolution of cell aggregation as the intracellular CheY concentration increases by induction with higher concentrations of arabinose. The scale bar is 10 μm. B. Quantitative characterization of intracellular c-di-GMP levels at different CheY concentrations. From top to bottom, they correspond to the ΔcheY strain (N=198), wildtype strain (N=249) and CheY overexpression strain (N=228), respectively.

To investigate whether a higher level of intracellular CheY-P increased the intracellular c-di-GMP level, we sought to detect the c-di-GMP level. We introduced the plasmid pCdrA-gfp into P. aeruginosa, using it as a c-di-GMP biosensor with the fluorescence intensity proportional to the c-di-GMP level45. To ensure the coexistence of the plasmids, the intracellular CheY concentration was controlled by the plasmid cheY-pME6032. The intracellular CheY levels induced by varying arabinose concentrations were measured, showing clear differences across the concentration gradient (Fig. S7). We found that the average single-cell fluorescence intensity of the CheY overexpression strain was 71.8% higher than that of the wild-type strain, while the c-di-GMP level of the cheY deletion strain was 58.6% lower than that of the wild-type (Fig. 4B).

The colocalization of both the signaling source (the kinase) and sink (the phosphatase) at the chemoreceptor complex at the cell pole results in a rapid decay of CheY-P concentration within approximately 0.2 µm from the cell pole, leading to a nearly uniform distribution elsewhere in the cell, as demonstrated by Vaknin and Berg46. This spatial arrangement effectively confines high CheY-P levels to the pole region. When the motor is also localized at the cell pole, the need for elevated CheY-P concentrations throughout the cytoplasm is reduced. Therefore, the colocalization of chemotaxis complex and flagellar motor facilitates the precise regulation of cell swimming direction within the low intracellular CheY-P concentration threshold, using locally distributed CheY-P molecules. This helps to avoid the occurrence of the aforementioned cross-pathway interference. To further confirm that CheY overexpression promotes aggregation through increased c-di-GMP levels, we performed additional experiments co-overexpressing CheY and a phosphodiesterase (PDE) from E. coli to reduce intracellular c-di-GMP. These experiments showed that PDE expression mitigates cell aggregation caused by CheY overexpression in P. aeruginosa (Fig. S8).

Summary and Discussion

P. aeruginosa harbors multiple signal transduction systems that regulate flagella-mediated swimming motility (Che pathway), pili-mediated interface twitching motility (Pil pathway), and biofilm formation (Wsp pathway)16. These systems allow it to thrive in various ecological niches within complex external environments. Here, employing a combination of chromosomal fluorescent protein fusion and flagellar labeling techniques in living cells, we directly observed that the chemotaxis complex of P. aeruginosa Che pathway and the flagellar motor share a high degree of consistency in their assembly sites. We found that the chemotaxis complex, comprising Che proteins, was consistently located at the cell pole where the flagellum was positioned. Based on these observations, we deduced the construction mode of the chemotactic network and flagellar motor during the entire cell growth cycle (Fig. 5A). As the cell body matures and elongates, a new flagellum grows at the opposing cell pole, accompanied by the assembly of a fresh chemotaxis complex. The cell then undergoes division from the middle, generating two daughter cells, each equipped with fully functional motility and chemotaxis capabilities. This ensures robust inheritance of these chemotaxis and motility-related macromolecular machines.

A. The construction mode of the chemotaxis network and flagellar motor throughout the complete cell growth cycle. B. The proximal growth mode of the P. aeruginosa flagellar motor and receptor clusters will effectively regulate the spatial range of CheY action, thus avoiding unintentional cross-pathway regulation.

Based on our understanding of the regulation mechanism of flagellar position in P. aeruginosa, we constructed a ΔflhF mutant and found that the position distribution of its chemotaxis complex is actively regulated by FlhF-mediated molecular mechanisms, rather than being passively restricted by membrane curvature factors as seen in E. coli. A similar active regulation phenomenon was previously reported in Vibrio cholerae, with the corresponding molecular regulation mechanism detailed extensively28. The flagella and chemotaxis complex of newborn V. cholerae cells are located at the old cell pole, with the polar anchor protein HubP distributed at both cell poles. As the cell grows, HubP recruits the ParA homologue ParC. Subsequently, ParC recruits ParP to assemble a new chemotaxis complex at the new cell pole. Upon cell division, HubP relocates to the middle of the cell, ensuring its robust presences at both poles of the daughter cells. Here, HubP recruits ParA homologues including FlhG, which influence flagellar motor assembly. However, previous experiments confirmed that the cell cycle-dependent polar localization of ParC and chemotactic proteins is independent of the flagellar position regulatory protein FlhF. This suggests that the localization of flagellar motors and chemotaxis complex in V. cholerae is controlled separately through different pathways. In contrast, flagellar motors were no longer robustly distributed at the cell pole after flhF was knocked out in P. aeruginosa, whereas the assembly site of chemotaxis complex and flagellar motor remained remarkably consistent. This suggests that FlhF regulates the distribution patterns of both and implies an interaction between them.

Removed in the revised version: We further constructed ΔfliG and ΔfliF mutants, which disrupted the assembly of the flagellar motor. Surprisingly, we observed that the assembly of chemotaxis complex was impaired under these conditions, with a significant decrease in the proportion of individuals exhibiting obvious Che protein fluorescent bright spots. In contrast, the chemotaxis complex assembly in the stator mutant (ΔmotAΔmotCD) was similar to that of the wild-type strain. It should be noted that P. aeruginosa has a four-tiered transcriptional regulatory circuit controlling flagellar biogenesis, in which FliG and FliF belong to class II and are regulated by the master transcriptional regulator FleQ47.

Earlier version: Chemotaxis complex-related genes (cheA, cheW, etc.) are regulated by intracellular free FliA proteins, and there is no evidence that FliA is regulated by proteins such as FliG. To further confirm that the expression of core chemotaxis proteins was not affected by the knockout of fliG, we performed Western blotting to measure CheY expression levels in wild-type, ΔfliF and ΔfliG strains, and found no significant differences (Fig. 3B). These results suggested that the integrity rather than functionality of the flagellar motor in P. aeruginosa is critical for the assembly of its chemotaxis complex, and its assembly efficiency is greatly reduced in the absence of flagellar motors. To eliminate the potential confounding effects of FlhF and flagellar motors, we constructed a ΔflhFΔfliF mutant strain. Under this condition, the proportion of individuals with obvious Che protein fluorescent bright spots dropped to ∼50%, indicating that passive regulation based on membrane curvature may still exist in P. aeruginosa, albeit to a lesser degree.

Revised version: We further constructed ΔfliG and ΔfliF mutants, which disrupted the assembly of the flagellar motor. Surprisingly, we observed that the assembly of chemotaxis complex was impaired under these conditions, with a significant decrease in the proportion of individuals exhibiting obvious Che protein fluorescent bright spots. In contrast, the chemotaxis complex assembly in the stator mutant (ΔmotAΔmotCD) was similar to that of the wild-type strain. It should be noted that P. aeruginosa has a four-tiered transcriptional regulatory circuit controlling flagellar biogenesis, in which FliG and FliF belong to class II and are regulated by the master transcriptional regulator FleQ47. Chemotaxis complex-related genes (cheA, cheW, etc.) are regulated by intracellular free FliA proteins. In E. coli, FliA is bound and inhibited by FlgM, and upon hook assembly, FlgM is secreted outside the cell, allowing FliA to trigger transcription of class III genes48,49, which include the chemosensory genes. This implies that if the hook is not assembled, late genes (including chemoreceptors) should not be expressed. However, Kaplan et al. reported that chemotaxis complexes were observed in Shewanella oneidensis ΔfliF mutant50, suggesting a different assembly process from that of E. coli. We therefore investigated whether knocking out fliF or fliG affects the expression level of che genes in P. aeruginosa. Since most experiments in this study focused on CheY, we performed Western blotting to measure CheY expression levels in wild-type, ΔfliF and ΔfliG strains, and found no significant differences (Fig. 3B). These results suggest that the observed decrease in the proportion of chemotaxis complexes after knocking out fliF and fliG should be attributed to the lack of motor assembly rather than changes in Che protein expression levels. To eliminate the potential confounding effects of FlhF and flagellar motors, we constructed a ΔflhFΔfliF mutant strain. Under this condition, the proportion of individuals with obvious Che protein fluorescent bright spots dropped to ∼50%, indicating that passive regulation based on membrane curvature may still exist in P. aeruginosa, albeit to a lesser degree.

Earlier version: Is the remarkable consistency in the spatial distribution of these two independent structural units an unintentional outcome or a hidden mystery? Given that Kulasekara et al. previously proposed that CheA and its phosphorylation state can cause c-di-GMP heterogeneity among individual P. aeruginosa cells through a specific PDE48,49, we sought to verify whether the phenotypic differences induced by CheY here are closely related to changes in the overall population level of c-di-GMP. We introduced a c-di-GMP monitoring system at the single-cell level and discovered that overexpression CheY substantially increases the intracellular c-di-GMP concentration. Previous studies generally believed that CheY-P, as a chemotactic regulatory protein, only influences flagellar motor switching. Here, we found that CheY-P, as a signaling molecule, can accomplish threshold-limited cross-pathway regulation. At low CheY-P concentrations, it is conservatively involved in motor switching regulation, as confirmed by previous studies. At high CheY-P concentrations, cell motility is suppressed, and c-di-GMP is triggered to express, though its specific mechanism needs further exploration. Unlike CheA, a structural protein of the chemotaxis complex, CheY is a protein that can shuttle in the cytoplasm and can also regulate the intracellular c-di-GMP level. This finding expands our understanding of the connection between the internal chemotaxis network and the c-di-GMP pathway. The intracellular environment is highly crowded50, and the proximal growth pattern shown in Fig. 5B has multiple physiological significances. On one hand, chemotaxis-related protein can precisely regulate cell motility with minimal synthesis costs. On the other hand, the chemotactic network and the flagellar motor are spatially integrated within the cell, ensuring their robust existence within the cell body and preventing mutual interference with collateral pathways.

Revised version: Is the remarkable consistency in the spatial distribution of these two independent structural units an unintentional outcome or a hidden mystery? Given that Kulasekara et al. previously proposed that CheA and its phosphorylation state can cause c-di-GMP heterogeneity among individual P. aeruginosa cells through a specific PDE51,52, we sought to verify whether the phenotypic differences induced by CheY here are closely related to changes in the overall population level of c-di-GMP. We introduced a c-di-GMP monitoring system at the single-cell level and discovered that overexpression CheY substantially increases the intracellular c-di-GMP concentration. Previous studies generally believed that CheY-P, as a chemotactic regulatory protein, only influences flagellar motor switching. Here, we found that CheY-P, as a signaling molecule, can accomplish threshold-limited cross-pathway regulation. At low CheY-P concentrations, it is conservatively involved in motor switching regulation, as confirmed by previous studies. At high CheY-P concentrations, cell motility is suppressed, and c-di-GMP is triggered to express, though its specific mechanism needs further exploration. Unlike CheA, a structural protein of the chemotaxis complex, CheY is a protein that can shuttle in the cytoplasm and can also regulate the intracellular c-di-GMP level. This finding expands our understanding of the connection between the internal chemotaxis network and the c-di-GMP pathway. The intracellular environment is highly crowded53, and the proximal growth pattern shown in Fig. 5B has multiple physiological significances. On one hand, chemotaxis-related protein can precisely regulate cell motility with minimal synthesis costs. On the other hand, the chemotactic network and the flagellar motor are spatially integrated within the cell, ensuring their robust existence within the cell body and preventing mutual interference with collateral pathways.

Many bacteria possess multiple signal transduction pathways. The orderly operation of each pathway within the micrometer-scale space presents a fascinating scientific problem. Here, we utilized the well-known chemotaxis network as a starting point to shed light on this problem, providing inspiration for future in-depth understanding of related issues.

Materials and methods

Strains and cell culture

Earlier version: The strains and plasmids used in this study are listed in Table 1. The Escherichia coli TOP10 strain was used for standard genetic manipulations. A single-colony isolate was grown in 3 ml of LB broth (1% Bacto tryptone, 0.5% yeast extract and 1% NaCl) overnight to saturation on a rotary shaker (250 rpm) at 37°C. An aliquot was diluted 1:100 into 10 ml of LB broth and grew to exponential phase. Appropriate antibiotics were added if necessary to prevent plasmid loss: for E. coli, 15 μg/ml Gentamicin and 25 μg/ml Tetracycline; for P. aeruginosa, 30 μg/ml Gentamicin and 50 μg/ml Tetracycline. To induce protein expression, Isopropyl β-D-1-thiogalactopyranoside (IPTG, 1.0 mM) was added to strains with pME6032-derivative vectors, and arabinose (0.005%, 0.01%, 0.05%, 0.1%) was added to strains with pJN105-derivative vectors. 2 ml of cells were harvested by centrifugation at 2000×g for 2 min, washed twice in an equal volume of motility buffer (MB) [50 mM potassium phosphate, 15 μM EDTA, 0.15 M NaCl, 5 mM Mg2+ and 10 mM lactic acid (pH 7.0)]51, and resuspended in 5 ml MB for subsequent fluorescence observation.

Revised version: The strains and plasmids used in this study are listed in Table 1. The Escherichia coli TOP10 strain was used for standard genetic manipulations. A single-colony isolate was grown in 3 ml of LB broth (1% Bacto tryptone, 0.5% yeast extract and 1% NaCl) overnight to saturation on a rotary shaker (250 rpm) at 37°C. An aliquot was diluted 1:100 into 10 ml of LB broth and grew to exponential phase. Appropriate antibiotics were added if necessary to prevent plasmid loss: for E. coli, 15 μg/ml Gentamicin and 25 μg/ml Tetracycline; for P. aeruginosa, 30 μg/ml Gentamicin and 50 μg/ml Tetracycline. To induce protein expression, Isopropyl β-D-1-thiogalactopyranoside (IPTG, 1.0 mM) was added to strains with pME6032-derivative vectors, and arabinose (0.005%, 0.01%, 0.05%, 0.1%) was added to strains with pJN105-derivative vectors. 2 ml of cells were harvested by centrifugation at 2000×g for 2 min, washed twice in an equal volume of motility buffer (MB) [50 mM potassium phosphate, 15 μM EDTA, 0.15 M NaCl, 5 mM Mg2+ and 10 mM lactic acid (pH 7.0)]54, and resuspended in 5 ml MB for subsequent fluorescence observation.

Construction of in-frame deletion mutants

Earlier version: We used polymerase chain reaction (PCR) to generate ∼1000 bp DNA fragments with upstream (Up) and downstream (Dn) sequences flanking the target gene (including flhF, fliF, motA and fliG) to be deleted. The Up and Dn DNA sequences were fused with the linearized pex18gm vector using Gibson assembly52. The resulting vectors were transformed into P. aeruginosa by electroporation, and the desired knockout mutants were obtained by double selections on gentamicin plates (LB plates with 30 µg/ml gentamicin) and sucrose plates (NaCl-free LB plates with 15% sucrose) at 37°C53.

Revised version: We used polymerase chain reaction (PCR) to generate ∼1000 bp DNA fragments with upstream (Up) and downstream (Dn) sequences flanking the target gene (including flhF, fliF, motA and fliG) to be deleted. The Up and Dn DNA sequences were fused with the linearized pex18gm vector using Gibson assembly55. The resulting vectors were transformed into P. aeruginosa by electroporation, and the desired knockout mutants were obtained by double selections on gentamicin plates (LB plates with 30 µg/ml gentamicin) and sucrose plates (NaCl-free LB plates with 15% sucrose) at 37°C56.

Generation of chromosomal fusions of yfp to P. aeruginosa gene

To generate eyfp fusion at the N-terminus of cheY, we first used PCR to generate upstream and downstream DNA sequences flanking the cheY gene. A subsequent PCR reaction aimed to yield the cheY-eyfp fusion with a 3× glycine linker (GGCGGAGGA), with pVS88 serving as the source of the enhanced yellow fluorescent (eyfp). The three DNA sequences, along with the linearized vector, were fused to cheY-eyfp-pex18gm using Gibson assembly. Finally, gene replacement was used to transfer the fusion constructs into the P. aeruginosa chromosome by homologous recombination.

Flagellar staining and fluorescence imaging

Earlier version: Flagellar filaments were labeled by following the protocol described previously54. Cells (1 ml of exponential-phase culture) were harvested by centrifugation at 2000×g for 10 min and washed twice in 1 ml MB. The final pellet was adjusted to a volume of ∼100 µl that concentrated the bacterial 10-fold. Alexa Fluor 568 maleimide (Invitrogen-Molecular Probes) was added to a final concentration of 20 μg/ml, and labeling was allowed to proceed for 30 min at room temperature with gyration at 80 rpm. Unused dye was then removed by washing the cells with MB three times, and the final pellet was resuspended in MB.

Revised version: Flagellar filaments were labeled by following the protocol described previously57. Cells (1 ml of exponential-phase culture) were harvested by centrifugation at 2000×g for 10 min and washed twice in 1 ml MB. The final pellet was adjusted to a volume of ∼100 µl that concentrated the bacterial 10-fold. Alexa Fluor 568 maleimide (Invitrogen-Molecular Probes) was added to a final concentration of 20 μg/ml, and labeling was allowed to proceed for 30 min at room temperature with gyration at 80 rpm. Unused dye was then removed by washing the cells with MB three times, and the final pellet was resuspended in MB.

For fluorescence imaging, 50 μl of cells were added to the chamber (constructed with two double-sticky tapes, a glass slide, and a poly-L-Lysine coated coverslip), incubated for 3 min, and rinsed with 100 μl MB. The boundary of the chamber was then sealed with Apiezon vacuum grease. The chamber was placed on a Nikon Ti-E inverted fluorescence microscope with a 100× oil-immersion objective and a sCMOS camera (Primer95B, Photometrics). The flagellar filament and chemotaxis complex of cells were observed separately with corresponding filter-set, using a 200-ms exposure time. To quantify the fluorescence intensity of the receptor cluster, we took the intensity maximum point as the center and extract a 3×3 pixel matrix around it. The mean value of the elements in this matrix represents the final fluorescence intensity.

Western Blotting

The total protein of P. aeruginosa strain cells in each treatment group was extracted and its concentration was measured with Bicinchoninic acid (BCA Protein Quantification Kit, Yeasen Biotechnology Co., Ltd) when the cells grew to exponential phase with same reduced turbidity (OD600 between 0.9 and 1.0). The exact amount of protein was subjected to SDS-PAGE electrophoresis, then transferred to a nitrocellulose membrane, which was blocked with 50 mg/mL skim milk in TBST buffer (20 mM Tris, 150 mM NaCl [pH7.4], 0.1%Tween 20) for 1h at room temperature. Since the N-terminal sequence of EYFP is almost identical to that of GFP, a rabbit anti-GFP antibody (1:2000, Yeasen Biotechnology Co., Ltd) and a Mouse anti-β-actin antibody (1:2000, Invitrogen) were added and incubated at 4°C overnight. The nitrocellulose membrane was gently washed with TBST for 20 min, three times. Subsequently, HRP labeled goat anti-rabbit IgG (1:4000, Abcam) and rabbit anti-mouse IgG (1:4000, Abcam) were added and incubated at room temperature for 1 h. The nitrocellulose membrane was again gently washed with TBST for 20 min, three times. Then, the membrane was developed using the ECL chemiluminescence detection system (Beyotime P0018FS) and visualized with the ChemiDoc XRS+ system (Bio-Rad). Three independent experiments were conducted for reproducibility.

Monitoring c-di-GMP signal at the single-cell level

For measurements of c-di-GMP signaling, the validated reporter plasmid pCdrA-gfp, which produces green fluorescent protein (GFP) in response to an increase in c-di-GMP, was introduced into the experimental system. A 1000× diluted culture was inoculated into the apparatus described in the previous section and allowed to adhere to the surface for ∼10 min, with the slight difference that the coverslips were not coated with poly-L-Lysine to avoid interference from background fluorescence. Bacteria were imaged using a Nikon Ti-E inverted microscope, which was equipped with an EMCCD camera (DU897, iXon3, Andor Technology). We used a 100× oil-immersion objective. Bacteria were illuminated with a 488-nm laser (Sapphire 488−200 mW, Coherent), using a standard GFP filter sets in the fluorescence imaging system with an exposure time of 500 ms. To ensure the consistency of culture and observation conditions, the data of wild-type and cheY-pME6032 carried strains were collected at the same time. More than 200 independent individuals of each strain were used for data analysis.

Single-cell c-di-GMP concentration was characterized by calculating fluorescence intensity. The average fluorescence intensity of a single cell was calculated by dividing the total fluorescence intensity of a single cell by the cell volume. Bacterial were simplified to a hollow cylindrical trunk with hollow hemispherical caps at both ends. The cell volume can be calculated by taking the length of the short axis of the cell body as the diameter of the sphere and the length of the long axis as the sum of the height of the cylinder and the diameter of the sphere. The mean fluorescence intensity of the cell was obtained by subtracting the average background fluorescence intensity from the total fluorescence intensity of the cell, and then dividing by the cell volume.